<div class="alert alert-success">

#### Lab 1
    
# Python Basics and Math Review

### EECS 245, Winter 2026 at the University of Michigan
    
</div>

### Instructions

Most labs will have Jupyter Notebooks, like this one, designed to supplement the in-person worksheet. 

To write and run code in this notebook, you have two options:

1. **Set up a Jupyter Notebook environment locally, and use `git` to clone our course repository (preferred)**. For instructions on how to do this, see the [**Environment Setup**](https://eecs245.org/env-setup) page of the course website. We have given you time in lab to follow these steps; if you have this notebook open, you probably already did them.
1. **Use the EECS 245 DataHub.** To do this, click the "code" link under Lab 1 on the course website. Log in with your uniqname and set a password.

To receive credit for the lab, you'll need to complete all 5 tasks in this notebook and show your lab TA. At the bottom of the notebook, we also demonstrate the workflow for submitting to Gradescope that you'll need to use for homeworks, though it's not necessary to submit to Gradescope for this lab.

## Overview: Jupyter Notebooks

---

### What is a Jupyter Notebook?

In the past, you may have been used to writing code in one window and running it in another (e.g. a Terminal).

<center><img src="imgs/text-editor-terminal.png" width=600></center>

**Jupyter Notebooks** allow us to write and run code within a single document. They also allow us to embed text and images and look at visualizations, which are useful in machine learning and data science workflows. (Jupyter stands for **Ju**lia, **Py**thon, and **R**, the three original languages they were designed to support.)

`.ipynb` is the extension for Jupyter Notebook files. `.ipynb` files can be opened and run in a few related applications, including JupyterLab, Jupyter Notebook, Jupyter Notebook Classic, and VSCode. If you're using the EECS 245 DataHub, the "code" link will open the standard Jupyter Notebook environment; if you're setting things up locally, the Tech Support page provides details on how to launch each of these interfaces.

### Cells

The **cell** is the basic building block of a Jupyter Notebook. There are two main types of cells.
- **Code cells** allow you to write and execute code. When run, code cells display the value of the last evaluated expression.
- **Markdown cells** allow you to write text and images that aren't Python code.
    - Markdown cells are always "run", except when you're editing them. 
    - Double-click this cell and see what happens!
    - Read more about Markdown [here](https://www.markdownguide.org/cheat-sheet/).
    
<center><img src='imgs/mdcell.png' width=90%>A code cell and Markdown cell, before and after being "run".</center>

### Using Python as a Calculator

To familarize ourselves with the notebook environment, let's run a few code cells involving arithmetic expressions. 

To run a code cell, **hit `shift` + `enter` (or `shift` + `return`) on your keyboard**. You could also hit the "▶️ Run" button in the toolbar at the top, but that is much, much slower, so don't get in the habit of doing that.

```python
# When you run this cell, the value of the expression appears, but isn't saved anywhere!
# These are comments, by the way.
17 ** 2
```

```python
# Integer division.
25 // 4
```

```python
min(-5.7, 1, 3) + max(4, 9, 7)
```

```python
# Why do we only see one line of output?
2 - 4
18 + 15.0
```

```python
# Strings can be created using single, double, or triple quotes.
# There's no difference between a string and a char.
'678' + "9" * 3
```

```python
'''November 26,
''' + "1998"
```

```python
# Put ? after the name of a function to see its documentation inline.
# All notebook interfaces support tab for autocompletion, too.
round?
```

## Variables and the Kernel

---

If you don't have any experience with Python, don't fret – we'll cover some of the basic features of the language here, and some additional features in the code component of Homework 1.

If you have used Python before, you might want to skip to the next section, where we introduce how [the autograder](#the-autograder) works.

### Python vs. C++ Variables

First, if you're familiar with C++ but not Python, here's a table summarizing some of their key differences.

| | Python | C++ |
| --- | --- | --- |
| Do I need to define<br>the type of a variable<br>beforehand? | <span style="color:orange"><b>No</b></span> <br>Python is dynamically typed. | <span style="color:blue"><b>Yes</b></span><br>C++ is statically typed. |
| Do I compile<br>my code before running it? | <span style="color:orange"><b>No</b></span><br>Python is interpreted;<br>Python code is converted to<br>bytecode line-by-line<br> at runtime.<br><small><br>In fact, the standard implementation<br> of Python is written in C (called CPython).</small> | <span style="color:blue"><b>Yes</b></span><br>The entirety of a<br> C++ program needs to be<br> compiled to bytecode<br> before it's run.<br><br><small>This is part of why C++ is <br>much <b>faster</b> than Python.</small> |

```python
# In C++, we'd need to define the type of count in advance, and it wouldn't be allowed to change.
# This works fine!
count = 7 + 9
count = "machine learning"
count
```

```python
type(count) # The type function returns the type of an object.
```

```python
# Type hints exist, but aren't enforced by default.
name: str = 'Junior'
name = 3.14
```

### Notebook Memory Model

Python may be new to you, but in addition, code in a Jupyter Notebook behaves a little differently than code in a text editor + Terminal setup.

We like to pretend your notebook has a brain 🧠. Everytime you run a cell with an assignment statement, it remembers that name-value binding. It will remember all name-value bindings as long as the current session is open, no matter how many cells you create or delete.

```python
# We defined this a few cells ago, but it still remembers the binding.
# This is a common pattern: writing the name of a variable in a cell of its own
# to check its value.
count
```

But, quitting your Terminal – or being inactive on DataHub for a few hours – ends your Jupyter Notebook session, and your notebook will forget everything it knows – **you’ll need to re-run all of your cells** the next time you open it.

With this in mind, you should aim to structure your code in a **reproducible** manner – so that others can trace your steps. Let's look at some practices you **should avoid ❌**. (By others, we mostly mean you, when you come back to work on your homework the next day.)

1. **Don't** delete cells that contain assignment statements.

```python
# To illustrate the issue, run this cell and then delete it.
age = 26
```

```python
# If the above cell has been run, this cell will run just fine, even if you 
# delete the cell above. However, once your notebook "forgets" all of 
# the variables it knows about, this cell will error, 
# since `age` won't be defined anywhere!
age + 15
```

2. **Don't** use a variable in a cell **above** where it is defined.

```python
# If you run the cell below first, then this cell will run just fine.
# However, once your notebook "forgets" all of the variables
# it knows about, and you run all of its cells in order,
# this will cause an error, because you are trying to use
# `weather` before its defined!
weather - 4
```

```python
# To illustrate the issue, run this cell FIRST, then the cell above.
weather = 72
```

3. **Don't** overwrite built-in names!

```python
min(2, 3)
```

```python
min = 17
```

```python
min(2, 3)
```

### Restarting the Kernel

If something doesn't seem right, you can **force** your notebook to forget everything it currently is remembering and give it a "fresh start".

To do so, first, save your notebook (by clicking the floppy disk icon or `CTRL/CMD + S`).

Then, **restart your kernel** by going to "Kernel" in the toolbar and clicking one of the options involving "Restart". Our preference is "Restart Kernel and Clear Outputs of All Cells" (or the equivalent in your UI), so that none of the previous cell outputs are visible. The kernel is like the engine of your notebook.

### Tip: Adding Cells

Often, to experiment, you'll want to create cells above or below the current one you're working in. The "+" button in the toolbar can help with this.

There are also keyboard shortcuts – first, hit `Escape` on your keyboard (to enter "Command Mode"), then hit the `A` key on your keyboard to add a cell above, or `B` to add a cell below. `DD` deletes a cell.

## The Autograder

---

Like you may have seen in other programming classes, your code in this class will be autograded – that is, automatically graded by the computer. The Python module we'll be using for autograding is called Otter Grader.

Run the cell below to import otter and initialize it for this notebook. In most homeworks, this cell will be the very first cell in your notebook.

```python
import otter
grader = otter.Notebook('lab01.ipynb')
```

Let's work through an example task to get a feel for how it works.

<div class="alert alert-info" markdown="1">

### Task 1: Seconds in an Hour

</div>
    
Below, you should see a question, a place to write your answer, and another cell containing `grader.check('task01')`. Running this last cell will check your answer to the question. If it's wrong, you'll see an error message. Try putting in a really small number, like 15, just to see what happens.

Assign `seconds_in_an_hour` to the number of seconds in an hour.

```python
seconds_in_an_hour = ...
seconds_in_an_hour
```

```python
grader.check("task01")
```

**In labs**, the tests you have access to in your notebook are always the exact same as the ones that we will run against your work on Gradescope.

**In homeworks**, this is not the case – homeworks have hidden test cases. In homeworks, the tests that you 
have access to in your notebook only verify that your answer is of the right data type and on the right track.

## Lists and Strings

---

Python has a variety of built-in data structures, including lists, dictionaries, sets, and tuples.

In this class, we'll most often use lists and dictionaries, along with more machine learning-specific data structures, like the `numpy` array, which we'll learn about next week.

A list is an **ordered** collection of values. To create a new list from scratch, we use [square brackets].

```python
mixed_list = [-2, 2.5, 'michigan', [1, 3], max] # Different types!
mixed_list
```

There are a variety of built-in functions that work with lists.

```python
max(['hey', 'hi', 'hello'])
```

The `append` **method** to add elements to the end of a list.

Since this is a method, we call it using "dot" notation, i.e. `groceries.append(...)` instead of `append(groceries, ...)`.

```python
groceries = ['eggs', 'milk']
groceries
```

```python
groceries.append('bread')
```

```python
groceries
```

**Important**: Note that `groceries.append('bread')` didn’t return anything, but groceries was modified.

We say `append` is **destructive**, because it does something other than return an output. We try to avoid destructive operations when possible.

```python
groceries + ['yogurt'] # This is a new list, not a modification of groceries!
```

Python, like most programming languages, is 0-indexed. This means that the **index**, or position, of the first element in a list is 0, not 1.<br><small>One reason: an element's index represents how far it is from the start of the list.</small>

```python
nums = [3, 1, 'dog', -9.5, 'michigan']
```

```python
nums[0]
```

```python
nums[3]
```

```python
nums[-1] # Counts from the end.
```

```python
nums[5]
```

We can use indexes to create a "slice" of a list. A slice is a new list containing elements from another list.

```python
nums
```

```python
nums[1:3]
```

In general,

```python
list_name[start : stop]
```

consists of all elements in list_name starting with index `start` and ending right before index `stop`.

```python
nums
```

```python
nums[0:4]
```

```python
# If you don't include 'start', the slice starts at the beginning of the list.
nums[:4]
```

```python
# If you don't include 'stop', the slice starts at the end of the list.
nums[-2:]
```

```python
# Interesting...
nums[::-1]
```

### Strings

Strings are similar to lists: they have indexes as well. Each element of a string can be thought of as a "character", which is a string of length 1.

```python
university = 'university of michigan'
```

```python
university[1]
```

```python
university[11:13]
```

```python
university[-8:]
```

Strings have various methods, but unlike `append`, they are not destructive – they return new strings.

```python
university.title()
```

```python
university.replace('i', 'I').split()
```

<div class="alert alert-info" markdown="1">

### Task 2: Uniqnames

</div>

The list `students` contains the (fake) @umich.edu email addresses of several students.

Using only the `students` variable, assign `third_last` to the **uniqname** (not email address) of the **third-last** student in the list. By "using only the `students` variable", we mean that you shouldn't hard-code the student's uniqname as a string just by looking at the list.

```python
students = [
    "jsmith42@umich.edu",
    "sarah.johnson@umich.edu",
    "m.chen@umich.edu",
    "alex.rodriguez@umich.edu",
    "kpatel23@umich.edu",
    "emma.wilson@umich.edu",
    "david1lee@umich.edu",
    "r.thompson@umich.edu",
    "nicole.garcia@umich.edu",
    "jbrown55@umich.edu",
    "tyler.anderson@umich.edu",
    "amy.kim@umich.edu"
]

third_last = ...
third_last
```

```python
grader.check("task02")
```

### Immutability

One key difference between lists and strings: you **can change** an element of a list, but **not** of a string.

If you want to change any part of a string, you must make a new string. This is because lists are **mutable**, while strings are **immutable**.

Below, before and after running `test_list[1] = 99`, `test_list` still refers to the same object in memory under the hood.

```python
test_list = [8, 0, 2, 4]
test_string = 'zebra'
```

```python
id(test_list) # Memory address of test_list.
```

```python
id(test_string)
```

```python
test_list[1] = 99
test_list
```

```python
id(test_list) # Same memory address!
```

```python
# This will error! Read the error message carefully.
test_string[1] = 'f'
```

```python
# Since we can't "change" test_string, we need to make a "new" string 
# containing the parts of it that we wanted.
# We can re-use the variable name test_string, though!
test_string = test_string[:1] + 'f' + test_string[2:]
test_string
```

_Most_ data structures – lists, dictionaries, `numpy` arrays, `pandas` DataFrames – are **mutable**, which means we need to be **extremely careful** when using them to modify them unexpectedly.

As an aside, Python Tutor, found at [pythontutor.com](https://pythontutor.com), allows you to visualize the execution of Python code.

Click [this link](https://pythontutor.com/render.html#code=test_list%20%3D%20%5B8,%200,%202,%204%5D%0Atest_string%20%3D%20'zebra'%0Atest_list%5B1%5D%20%3D%2099%0Atest_string%5B1%5D%20%3D%20'f'&cumulative=false&curInstr=0&heapPrimitives=nevernest&mode=display&origin=opt-frontend.js&py=311&rawInputLstJSON=%5B%5D&textReferences=false) to visualize the previous slide's code, or run the cell below to see it embedded in this notebook.

```python
from IPython.display import IFrame
def test_pt_example():
    src = "https://pythontutor.com/iframe-embed.html#code=test_list%20%3D%20%5B8,%200,%202,%204%5D%0Atest_string%20%3D%20'zebra'%0Atest_list%5B1%5D%20%3D%2099%0Atest_string%5B1%5D%20%3D%20'f'&codeDivHeight=400&codeDivWidth=350&cumulative=false&curInstr=-1&heapPrimitives=nevernest&origin=opt-frontend.js&py=311&rawInputLstJSON=%5B%5D&textReferences=false"
    width = 800
    height = 500
    display(IFrame(src, width, height))

test_pt_example()
```

Python is notoriously opaque when it comes to variables and pointers. Here's a [good reference](https://nedbatchelder.com/text/names.html).

## Functions and `if`-Statements

---

### Indentation

In C++, to define code blocks, you used `{`curly brackets`}`.

```cpp
            double future_value(double present_value, double APR, int months) {
                double r = APR / 12.0 / 100.0;
                return present_value * pow(1 + r, months);
            }
```

In Python, you use a colon`:` and then **indent** the following lines by either a **tab** or **four spaces**.

```python
def future_value(present_value, APR, months):
    r = APR / 12 / 100
    return present_value * (1 + r) ** months
```

```python
future_value(100, 7, 36)
```

The `def` keyword defines a new function. `if`-statements, `for`-loops, and `while`-loops work similarly as in other languages.

The remainder of the lab contains practice with functions and `if`-statements; you'll practice writing loops in Homework 1.

<div class="alert alert-info" markdown="1">

### Task 3: Tax Brackets

</div>

The United States, like many countries, uses a progressive tax bracket system. This means that as your earnings increase, the percentage of your earnings you owe in tax also increases. In addition, the US tax system uses marginal tax brackets – what this means is that US taxpayers pay different tax percentages on different "chunks" of their earnings.

Let's suppose the tax brackets for the 2024 tax year for single filers are defined by the table below. These are close to the actual brackets, but for simplicity's sake we'll use 5 brackets instead of 7. 

| Tax Bracket | Taxable Income |
| --- | --- |
| 10% | [$0, \\$11,000] |
| 12% | (\\$11,000, \\$44,725] |
| 22% | (\\$44,725, \\$95,375] |
| 24% | (\\$95,375, \\$182,100] |
| 32% | Over \\$182,100 |

The notation $(a, b]$ means "greater than $a$ and less than or equal to $b$". For example, someone with a taxable income of \\$44,725 is in the 12% bracket, but someone with a taxable income of \\$44,725.01 is in the 22% bracket. For simplicity, we will only test your code on integer values of taxable income.

If someone has a taxable income of \\$75,000, we say they are in the 22% tax bracket. However, such an individual doesn't owe 22% of \\$75,000 in taxes. Instead, they owe:
- 10% of \\$11,000, **plus**
- 12% of \\$33,725 (which is \\$44,725 - \\$11,000), **plus**
- 22% of \\$30,275 (which is \\$75,000 - \\$44,725).

Complete the implementation of the function `tax_bracket`, which takes in a taxable income, `income`, and returns the tax bracket it is in, as a **proportion**. Assume that `income` is a non-negative integer. Example behavior is given below.

```python
>>> tax_bracket(75000)
0.22

>>> tax_bracket(402150)
0.32
```

```python
def tax_bracket(income):
    ...

# Feel free to change this input to make sure your function works correctly.
tax_bracket(75000)
```

```python
grader.check("task03")
```

<div class="alert alert-info" markdown="1">

### Task 4: Running, Eating, Living

</div>

Complete the implementation of the function `is_ing`, which takes in a string `s` and returns `True` if `s` ends in "ing" and `False` otherwise. Example behavior is given below.

```python
>>> is_ing('running')
True

# Not case sensitive!
>>> is_ing('eatInG')
True

>>> is_ing('science')
False

>>> is_ing('25')
False
```

```python
def is_ing(s):
    ...

# Feel free to change this input to make sure your function works correctly.
is_ing('running')
```

```python
grader.check("task04")
```

<div class="alert alert-info" markdown="1">

### Task 5: Pump the Brakes

</div>


Complete the implementation of the function `pump`, which takes in a string `base` and integer `rep` and returns a string with the same first and last characters as `base`, and with everything except the first and last characters in `base` repeated `rep` times. Example behavior is given below.

```python
>>> pump('zebra', 4)
'zebrebrebrebra'

# If base is <= 2 characters long, return base.
>>> pump('hi', 2)
'hi'
```

_Hint: You don't need a loop to answer this question. Instead, experiment with what happens when you multiply a string by an integer!_

```python
def pump(base, rep):
    ...

# Feel free to change this input to make sure your function works correctly.
pump('zebra', 4)
```

```python
grader.check("task05")
```

## Finish Line 🏁

Congratulations! You're ready to submit Lab 1.

**To get credit for the lab, all you need to do is show your lab TA that you've finished this notebook.**

---

Optionally, in addition, you can practice the Gradescope submission workflow, which you'll need to do in Homework 1. If you want to go this route, follow these steps:

1. Select `Kernel -> Restart & Run All` to ensure that you have executed all cells, including the test cells. **There will be cells that error from the "Notebook Memory Model" section; comment the buggy lines of code out.**

2. Read through the notebook to make sure everything is fine and all tests passed.
3. Run the cell below to run all tests, and make sure that they all pass.
4. Download your notebook using `File -> Download`, then upload your notebook to Gradescope under "Lab 1".
5. Stick around for a few minutes while the Gradescope autograder grades your work. Make sure you see that all **public tests** have passed on Gradescope.
6. Show your lab TA once you've finished your worksheet too.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task01 (points: None)
```
>>> isinstance(seconds_in_an_hour, int)
True
```
```
>>> seconds_in_an_hour > 3000
True
```
failure_message: Your answer is too small!
```
>>> seconds_in_an_hour == 3600
True
```
success_message: Nice job!


Tests for task02 (points: None)
```
>>> isinstance(third_last, str)
True
```
```
>>> '@' not in third_last
True
```
failure_message: Make sure your answer doesn't contain @umich.edu.
```
>>> third_last == 'jbrown'
False
```


Tests for task03 (points: None)
```
>>> tax_bracket(11000) == 0.10
True
```
```
>>> tax_bracket(100_000) == 0.24
True
```
```
>>> tax_bracket(1) == 0.10
True
```
```
>>> tax_bracket(13_332) == 0.12
True
```
```
>>> tax_bracket(999_999) == 0.32
True
```


Tests for task04 (points: None)
```
>>> is_ing('running')
True
```
```
>>> is_ing('eatInG')
True
```
```
>>> not is_ing('science')
True
```
```
>>> not is_ing('25')
True
```
```
>>> not is_ing('')
True
```
```
>>> is_ing('ING')
True
```


Tests for task05 (points: None)
```
>>> pump('zebra', 4) == 'zebrebrebrebra'
True
```
```
>>> pump('hi', 2) == 'hi'
True
```
```
>>> pump('', 74) == ''
True
```
```
>>> pump('hello', 1) == 'hello'
True
```
```
>>> pump('zebra', 6) == 'zebrebrebrebrebrebra'
True
```
```
>>> pump('wolverine', 5) == 'wolverinolverinolverinolverinolverine'
True
```
```
>>> pump('michigan', 15).count('ichiga') == 15
True
```
